## CSCI 470 Activities and Case Studies

1. For all activities, you are allowed to collaborate with a partner. 
1. For case studies, you should work individually and are **not** allowed to collaborate.

By filling out this notebook and submitting it, you acknowledge that you are aware of the above policies and are agreeing to comply with them.

Some considerations with regard to how these notebooks will be graded:

1. Cells in which "# YOUR CODE HERE" is found are the cells where your graded code should be written.
2. In order to test out or debug your code you may also create notebook cells or edit existing notebook cells other than "# YOUR CODE HERE". We actually highly recommend you do so to gain a better understanding of what is happening. However, during grading, **these changes are ignored**. 
2. You must ensure that all your code for the particular task is available in the cells that say "# YOUR CODE HERE"
3. Every cell that says "# YOUR CODE HERE" is followed by a "raise NotImplementedError". You need to remove that line. During grading, if an error occurs then you will not receive points for your work in that section.
4. If your code passes the "assert" statements, then no output will result. If your code fails the "assert" statements, you will get an "AssertionError". Getting an assertion error means you will not receive points for that particular task.
5. If you edit the "assert" statements to make your code pass, they will still fail when they are graded since the "assert" statements will revert to the original. Make sure you don't edit the assert statements.
6. We may sometimes have "hidden" tests for grading. This means that passing the visible "assert" statements is not sufficient. The "assert" statements are there as a guide but you need to make sure you understand what you're required to do and ensure that you are doing it correctly. Passing the visible tests is necessary but not sufficient to get the grade for that cell.
7. When you are asked to define a function, make sure you **don't** use any variables outside of the parameters passed to the function. You can think of the parameters being passed to the function as a hint. Make sure you're using all of those variables.
8. Finally, **make sure you run "Kernel > Restart and Run All"** and pass all the asserts before submitting. If you don't restart the kernel, there may be some code that you ran and deleted that is still being used and that was why your asserts were passing.

# Data Visualization

## What is a data visualization?

> A visualization is any kind of visual representation of information designed to enable communication, analysis, discovery, exploration, etc.

> A data visualization is a display of data designed to enable analysis, exploration, and discovery.


Cairo, Alberto (2016-02-08). The Truthful Art: Data, Charts, and Maps for Communication

## Importance of data visualizations

Most people think that descriptive statistics are sufficient to fully understand your data. Unfortunately this is not true. Descriptive statistics can be very misleading and data visualizations can help us uncover some insights.

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
%matplotlib inline
plt.style.use("ggplot")
```

The data below is called Anscombe's quartet. It contains 4 different datasets (one, two, three, four) and each one has x and y values.

```python
one_x = [10.0, 8.0, 13.0, 9.0, 11.0, 14.0, 6.0, 4.0, 12.0, 7.0, 5.0]
one_y = [8.04, 6.95, 7.58, 8.81, 8.33, 9.96, 7.24, 4.26, 10.84, 4.82, 5.68]
two_x = [10.0, 8.0, 13.0, 9.0, 11.0, 14.0, 6.0, 4.0, 12.0, 7.0, 5.0]
two_y = [9.14, 8.14, 8.74, 8.77, 9.26, 8.10, 6.13, 3.10, 9.13, 7.26, 4.74]
three_x = [10.0, 8.0, 13.0, 9.0, 11.0, 14.0, 6.0, 4.0, 12.0, 7.0, 5.0]
three_y = [7.46, 6.77, 12.74, 7.11, 7.81, 8.84, 6.08, 5.39, 8.15, 6.42, 5.73]
four_x = [8.0, 8.0, 8.0, 8.0, 8.0, 8.0, 8.0, 19.0, 8.0, 8.0, 8.0]
four_y = [6.58, 5.76, 7.71, 8.84, 8.47, 7.04, 5.25, 12.50, 5.56, 7.91, 6.89]

xs = [one_x, two_x, three_x, four_x]
ys = [one_y, two_y, three_y, four_y]
```

Your first task is to load the data from Anscombe's quartet into a pandas [DataFrame](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.html). This [guide](https://pandas.pydata.org/pandas-docs/stable/dsintro.html#dataframe) may be helpful if you are not already familiar with pandas.

<!-- grade_id: cell-2b46fab363c326f0, points: 0, grade: false, solution: true -->
```python
# Save the values into a dataframe called `df` with column names "x1", "x2", "x3", "x4", "y1", "y2", "y3", "y4"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3bfc8754b6cb3fcd, points: 10, grade: true, solution: false -->
```python
assert isinstance(df, pd.DataFrame)
assert all(col in df.columns for col in ["x1", "x2", "x3", "x4", "y1", "y2", "y3", "y4"])
assert len(df) == 11
```

In this next section, find the mean and variance of each series and compare the x's and y's. You may want to use the [Series.mean()](https://pandas.pydata.org/pandas-docs/version/0.22/generated/pandas.Series.mean.html) and [Series.var()](https://pandas.pydata.org/pandas-docs/version/0.22/generated/pandas.Series.var.html) or [DataFrame.mean()](https://pandas.pydata.org/pandas-docs/version/0.22/generated/pandas.DataFrame.mean.html) and [DataFrame.var()](https://pandas.pydata.org/pandas-docs/version/0.22/generated/pandas.DataFrame.var.html) functions

<!-- grade_id: cell-34448fbcbc07fc81, points: 0, grade: false, solution: true -->
```python
# Set the value of xmeans to a list of the mean of each of ["x1", "x2", "x3", "x4"]

# YOUR CODE HERE
raise NotImplementedError()


for xmean, xSeries in zip(xmeans, ["x1", "x2", "x3", "x4"]):
    print(f"The mean of {xSeries} is {xmean}.")
```

<!-- grade_id: cell-e04b9f460d75aad8, points: 10, grade: true, solution: false -->
```python
np.testing.assert_allclose(np.array(xmeans), 9) 
```

<!-- grade_id: ymeans, points: 0, grade: false, solution: true -->
```python
# Set the value of ymeans to a list of the mean of each of ["y1", "y2", "y3", "y4"]

# YOUR CODE HERE
raise NotImplementedError()


for ymean, ySeries in zip(ymeans, ["y1", "y2", "y3", "y4"]):
    print(f"The mean of {ySeries} is {ymean}.")
```

<!-- grade_id: ymeans-test, points: 10, grade: true, solution: false -->
```python
np.testing.assert_allclose(np.array(ymeans), 7.5, atol=0.1) 
```

<!-- grade_id: cell-e9e7895bdf3ae614, points: 0, grade: false, solution: true -->
```python
# Set the value of xvars to a list of the variance of each of ["x1", "x2", "x3", "x4"]

# YOUR CODE HERE
raise NotImplementedError()


for xvar, xSeries in zip(xvars, ["x1", "x2", "x3", "x4"]):
    print(f"The mean of {xSeries} is {xvar}.")
```

<!-- grade_id: cell-ae0d5b4b1d2e02f3, points: 10, grade: true, solution: false -->
```python
np.testing.assert_allclose(np.array(xvars), 11) 
```

<!-- grade_id: cell-61d16404c42f97f1, points: 0, grade: false, solution: true -->
```python
# Set the value of yvars to a list of the variance of each of ["y1", "y2", "y3", "y4"]

# YOUR CODE HERE
raise NotImplementedError()


for yvar, ySeries in zip(yvars, ["y1", "y2", "y3", "y4"]):
    print(f"The mean of {ySeries} is {yvar}.")
```

<!-- grade_id: cell-664ddc1141e9165a, points: 10, grade: true, solution: false -->
```python
np.testing.assert_allclose(np.array(yvars), 4.12, atol=0.01) 
```

In this next section, find the correlation and best fit line of each series and compare them. 

For correlation, you may find [Dataframe.corr](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.corr.html) to be helpful. Also see [this stack overflow answer](https://stackoverflow.com/a/42580124/818687).

For finding the best fit line, you may use [np.polyfit](https://docs.scipy.org/doc/numpy/reference/generated/numpy.polyfit.html).

<!-- grade_id: cell-2a970ac465d6e52d, points: 0, grade: false, solution: true -->
```python
# Set the value of corrs to be a list of correlations between x and y for each series
# Ex. the 0th index should be the correlation between x1 and y1

# YOUR CODE HERE
raise NotImplementedError()


for c, xSeries, ySeries in zip(corrs, ["x1", "x2", "x3", "x4"], ["y1", "y2", "y3", "y4"]):
    print(f"The correlation between {xSeries} and {ySeries} is {c}.")
```

<!-- grade_id: cell-2bc85947297ef229, points: 10, grade: true, solution: false -->
```python
np.testing.assert_allclose(np.array(corrs), 0.816, atol=0.01) 
```

<!-- grade_id: cell-e5dbbc03d0d125f4, points: 0, grade: false, solution: true -->
```python
# Determine the coefficients for each series and save them to the coeffs variable
# Find the best fit line

# YOUR CODE HERE
raise NotImplementedError()

for i, (m,b) in enumerate(coeffs):
    print(f"The fitted line for series {i+1} is y = {m}x + {b}")
```

<!-- grade_id: cell-e50306da72dca59c, points: 10, grade: true, solution: false -->
```python
np.testing.assert_allclose(np.array(coeffs), np.array([[0.5, 3.0], [0.5, 3.0], [0.5, 3.0], [0.5, 3.0]]), atol=0.01) 
```

From the above descriptive analytics, we would imagine all 3 datasets to be very similar. Let's plot them and see the patterns they exhibit.

```python
buffer = 3
maxX = max(max(xs[0]),max(xs[1]),max(xs[2]),max(xs[3]))
maxY = max(max(ys[0]),max(ys[1]),max(ys[2]),max(ys[3]))

x = np.arange(0,maxX+buffer)

fig, pltAxes = plt.subplots(ncols=2, nrows=2, figsize=(12,8))
axes = pltAxes.ravel()

for i, ax in enumerate(axes):
    fx = coeffs[i][0]* x + coeffs[i][1]
    ax.plot(xs[i],ys[i],"bo")              # Plot the x and y values
    ax.set_ylim(0, maxY + buffer)          # Use the same scale
    ax.set_xlim(0, maxX + buffer)
    ax.plot(x,fx,"r-")                      # Plot the line
    ax.set_title("Series {}".format(i+1), fontsize=12)   # Add the title for the dataset

fig.suptitle("Plotting Anscombe's Quartet", fontsize=14)
fig.subplots_adjust(hspace=.5)

plt.show()
```

**We can see from the above plots that the data from each series is actually quite different.**

- Series 1 exhibits a linear trend with some gaussian noise
- Series 2 exhibits a parabolic trend
- Series 3 exhibits a linear trend but has an outlier in the dataset
- Series 4 has one outlier and the remaining points all in various y values for similar x values

This shows us the importance of plotting data. Descriptive statistics are important, but they are not sufficient.

## Visual Encoding

We are able to perceive information with varying accuracy and generality based on how it is encoded. The following list is of a variety of ways to encode information and it is sorted from most accurate to most general.


*Most accurate*
1. Points on a scale
1. Length
1. Direction / Angle
1. Area
1. Volume
1. Shading and saturation
1. Color and hue

*Most general*


## Plotting

In this section we'll go over generating some random data and plotting them.

We will use the pandas [DataFrame.plot](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.plot.html) function as well as a variety of functions from [matplotlib](https://matplotlib.org/tutorials/index.html). This [visualization guide](https://pandas.pydata.org/pandas-docs/stable/visualization.html) from pandas is a great resource as well.


There are many ways to visualize data. The following are commonly used plots and when to use them:

1. Scatter plots - investigate the relationship between two variables
1. Bar/column charts - compare values of multiple samples
1. Histograms - investigate the distribution of the data
1. Boxplots - compare distributions of variables
1. Line charts - invetigate trends over time 
1. Pie charts - compare proportions (**You are usually better off using a different chart such as a bar chart**)


### Scatter Plots

Investigate the relationship between two variables

```python
# A simple example of a scatter plot
x = range(0,11,2)
y = range(0,17,3) + np.random.rand(6)*2
plt.plot(x,y, "bo")
plt.xlabel("x")
plt.ylabel("y")
plt.title("x vs y")
plt.show()
```

With scatter plots, we often want a reference line. This could be a best fit line or a mean/ specific percentile.

For best fit line, consider using [poly1d](https://docs.scipy.org/doc/numpy-1.14.0/reference/generated/numpy.poly1d.html#numpy.poly1d)

<!-- grade_id: cell-b22b71b5e1906329, points: 10, grade: true, solution: true -->
```python
# Determine the best fit line and set the function or poly1d object to line

# YOUR CODE HERE
raise NotImplementedError()

plt.plot(x,y,"bo")
plt.plot(x, line(x))
plt.xlabel("x")
plt.ylabel("y")
plt.title("x vs y")
plt.show()
```

<!-- grade_id: cell-224d3300d6c7b42c, points: 10, grade: true, solution: true -->
```python
# Determine the average line and set the values to avg

# YOUR CODE HERE
raise NotImplementedError()

plt.plot(x,y,"bo")
plt.plot(x, avg)
plt.xlabel("x")
plt.ylabel("y")
plt.title("x vs y")
plt.show()
```

Now let's create some data with some correlations and look at how we can handle having more than just 2 features.

In this example, we'll create 5 features.

```python
np.random.seed(0)

a = 20 * np.random.normal(0, 1, 100) + 12
b = a + 5 * np.random.rand(100) # Create a positive correlation with a
c = a * 2 * np.random.normal(0,1,100) - (b * 4 * np.random.normal(0,1,100)) # Create a relationship with both a and b
d = 45 * np.random.rand(100) 
e = -0.7* d -  10 * np.random.rand(100) # Create a negative correlation

df = pd.DataFrame({"a":a, "b": b, "c":c, "d":d, "e": e})
```

```python
pd.plotting.scatter_matrix(df, figsize=(10, 10))
plt.show()
```

```python
# We can set the diagonal to be a density line instead of a histogram as well
pd.plotting.scatter_matrix(df, figsize=(10, 10), diagonal="kde")
plt.show()
```

Scatter matrixes are extremely useful when investigating data from more than 2 features. Feel free to modify the relationships between the 5 features and take a look at the resulting output. This will help you understand patterns when you see them in real world datasets.

The [random module](https://docs.scipy.org/doc/numpy/reference/routines.random.html#distributions) from numpy has a variety of distributions that you can try out.

### Histograms

Use histograms to understand the distribution of a data set.

```python
uniforms = [np.random.rand() for i in range(1000)]
plt.hist(uniforms)
plt.title("Generated Uniforms")
plt.show()
```

```python
normals = [np.random.normal() for i in range(1000)]
plt.hist(normals)
plt.title("Generated Normals")
plt.show()
```

```python
plt.hist(normals,100)
plt.title("Generated Normals")
plt.show()
```

```python
plt.hist(normals,100, density=True)
plt.title("Normalized Generated Normals")
plt.show()
# Note the y axis!
```

Selecting the number of bins for a histogram can lead to varying insights.


>Histograms are fantastic, but make sure your bin-width/number is chosen well. This is the _exact_ same data, plotted with different bin-widths. Notice that the pattern doesn't necessarily get clearer as bin num increases.
> - [Nick Strayer](https://twitter.com/NicholasStrayer)
> [Twitter thread](https://twitter.com/NicholasStrayer/status/1026893778404225024)

Check out the [demo](http://nickstrayer.me/histogram_bins/)



There are proper ways to select the number of bins in a histogram. Some are outlined in [this guide](http://docs.astropy.org/en/stable/visualization/histogram.html) of a python library that you may use.

### Bar Charts

Bar charts are used to compare data.

**Always start the feature scale at 0.**

```python
# We'll generate some random sales data for 3 locations and 4 products
locations = ["Denver", "Golden", "Boulder"]
n = len(locations)
np.random.seed(0)
appleSales = [np.random.randint(3000,5000) for i in range(n)]
orangeSales = [np.random.randint(2000,5000) for i in range(n)]
avocadoSales = [np.random.randint(4000,5000) for i in range(n)]
cherrySales = [np.random.randint(2500,5000) for i in range(n)]
simpleSales = pd.DataFrame({
    "Location": locations, "Apples": appleSales, "Oranges": orangeSales, "Avocados": avocadoSales, "Cherries": cherrySales
})
```

Here we'll start using the [DataFrame.plot method](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.plot.html).

```python
simpleSales[["Location", "Apples"]].plot(kind="bar", x="Location", legend=False, title="Apple Sales")
plt.show()
```

Column (vertical bar) charts are great but Bar charts (horizontal) make the labels much easier to read. You can see that in the following chart:

```python
simpleSales[["Location", "Apples"]].plot(kind="barh", x="Location", title="Apple Sales", legend=None)
plt.show()
```

Bar charts are great for comparing samples. One way to make them more effective and easy to read is to sort the data's values. The default sorting isn't always useful so it's important to explicitly specify the sorting you'd want. You can do so by sorting the dataframe itself.

```python
simpleSales.sort_values("Apples", ascending=False)[["Location", "Apples"]].plot(kind="barh", 
                                                                                x="Location", 
                                                                                title="Apple Sales", 
                                                                                legend=None)
plt.show()
```

Since we are perceiving the relative distance on a line, it is important that bar charts always start the scale at 0. We create very misleading charts by changing the minimum limit.

**Bar charts should always start at 0**

```python
simpleSales.sort_values("Apples", ascending=False)[["Location", "Apples"]].plot(kind="barh", 
                                                                                x="Location", 
                                                                                title="Misleading Apple Sales", 
                                                                                legend=None)
plt.xlim(3500, 5000)
plt.show()
```

* To compare multiple series, use side by side bars/columns instead of stacked bar charts. 
* Stacked bar charts make it difficult to compare values since they don't start at the same level.

```python
simpleSales.sort_values("Apples", ascending=False).plot(kind="barh", x="Location")
plt.title("Sales by Location")
plt.tight_layout()
plt.legend(loc='upper center', bbox_to_anchor=(0.5, -.1),ncol=4)
plt.show()
```

## Boxplots

Use boxplots to compare distributions.

```python
simpleSales[["Apples", "Oranges", "Avocados", "Cherries"]].plot(kind="box", title="Sales")
plt.ylim(0,5500)
plt.title("Sales")
plt.show()
```

In a boxplot, the line inside the boxplot represents the median and the box itself represents 50% of the data (showing the 25% and 75% quartiles). The lines at the end, outside the box show the minimum and maximum value.

A great illustration of this is shown below:

![https://www.wellbeingatschool.org.nz/information-sheet/understanding-and-interpreting-box-plots](https://www.wellbeingatschool.org.nz/sites/default/files/W@S_boxplot-labels.png)

Similar to bar charts, we can also create horizontal boxplots.

```python
simpleSales[["Apples", "Oranges", "Avocados", "Cherries"]].plot(kind="box", title="Sales", vert=False)
plt.xlim(0,5500)
plt.title("Sales")
plt.show()
```

### Line Charts

Line charts are used to plot continuous data. That is data about the same entity that changes over time.

```python
np.random.seed(0)
nMonths = 12
sales = [np.random.randint(2000,5000) for i in range(nMonths)]
months = [str(i) for i in range(1,nMonths+1)]
```

```python
plt.plot(sales)
plt.xticks(range(len(sales)),months)
plt.title("Sales over time")
plt.xlabel("Month")
plt.ylabel("Sales")
# plt.ylim(0, 6000) # Line charts do not have to start at 0 but it is often useful that they do
plt.show()
```

Graphing multiple lines on the same chart may be hard to follow.

```python
year_1 = [np.random.randint(3000,5000) for i in range(12)]
year_2 = [np.random.randint(4000,5000) for i in range(12)]
year_3 = [np.random.randint(2000,5000) for i in range(12)]
```

```python
plt.plot(year_1, label="Year 1")
plt.plot(year_2, label="Year 2")
plt.plot(year_3, label="Year 3")
plt.xticks(range(len(year_1)),months)
plt.legend(loc="best")
plt.title("Sales over time")
plt.xlabel("Month")
plt.ylabel("Sales")
# plt.ylim(0, 6000)
plt.show()
```

Instead, use multiple panels of line charts using the same scale. You can also add a line to show the average for all the data sets as a reference point.

```python
avg_year_1 = [np.mean(year_1)] * 12
avg_year_2 = [np.mean(year_2)] * 12
avg_year_3 = [np.mean(year_3)] * 12
avg_total = [np.mean([year_1,year_2, year_3])] * 12

buffer = 500
maxY = max(max(year_1), max(year_2), max(year_3))
minY = min(min(year_1), min(year_2), min(year_3))

monthRange = range(1,13)

monthlyMean = np.mean([year_1,year_2, year_3], axis=0)

years = [year_1, year_2, year_3]
averages = [avg_year_1, avg_year_2, avg_year_3]
colors = ["r-", "b-", "g-"]

fig, axes = plt.subplots(nrows=3)

for i, ax in enumerate(axes):
    ax.plot(monthRange ,years[i], colors[i])
    ax.set_ylim(minY-buffer, maxY + buffer)
    ax.set_xlim(1, 12)
    ax.set_title("Year {}".format(i+1), fontsize=6)
#     ax.plot(monthRange, avg_year_1, "r--")
#     ax.plot(monthRange, avg_year_2, "b--")
#     ax.plot(monthRange, avg_year_3, "g--")
#     ax.plot(monthRange, avg_total, "k--")
#     ax.plot(monthRange, monthlyMean, "k--")


fig.suptitle("Comparing yearly performance", fontsize=9)
fig.subplots_adjust(hspace=.5)
plt.show()
```

Having the data appear in panels allows us to see the trend of each year and also compare how sales were across years for a given month.

### Pie Charts

Pie charts can be used to show proportion of data. You shouldn't use pie charts as there are usually better ways to display the same information. You can read more about why [here](http://www.businessinsider.com/pie-charts-are-the-worst-2013-6).

```python
percentSales = pd.Series(year_1, index=range(1,13), name="Sales")

percentSales.sort_values(ascending=True).plot(kind="pie")
plt.show()
```

Even though the data in the pie chart is sorted, it is still difficult to determine which month had the highest percent sales.

In this section, plot the same data using a bar chart instead of a pie chart. Make sure you sort the values when presenting the information.

<!-- grade_id: cell-cc2fd888aa5a24ff, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
print("The month with the highest sales is...")
```

From your chart, you should be able to easily determine the month with the highest sales.

## Choosing a visualization

Still not sure which visualization to use?

Use [this flow chart](https://extremepresentation.typepad.com/files/choosing-a-good-chart-09.pdf)

## Feedback

<!-- grade_id: cell-a96e2ab81584b833, points: 0, grade: false, solution: true -->
```python
def feedback():
    """Provide feedback on the contents of this exercise
    
    Returns:
        string
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a61433551b1618ef, points: 0, grade: true, solution: false -->
```python

```